Modulation of RNA processing genes during sleep-dependent memory

Author response:

The following is the authors’ response to the previous reviews.

eLife Assessment

The aim of this valuable study is to identify novel genes involved in sleep regulation and memory consolidation. It combines transcriptomic approaches following memory induction with measurements of sleep and memory to discover molecular pathways underlying these interlinked behaviors. The authors explore transcriptional changes in specific mushroom body neurons and suggest roles for two genes involved in RNA processing, Polr1F and Regnase-1, in the regulation of sleep and memory. Although this work exploits convincing and validated methodology, the strength of the evidence is incomplete to support the main claim that these two genes establish a definitive link between sleep and memory consolidation.

We appreciate the reconsideration of our manuscript and recognize that we should have toned down the claims, especially with respect to the link between sleep and memory consolidation. We have now changed the title, the abstract and the main text and also Figure 5 to essentially just state our findings. While there is a little speculation in the Discussion, we point out that future work would be required to draw conclusions. We believe the manuscript still represents a considerable advance in showing the modulation of RNA processing genes during sleep-dependent memory consolidation in the relevant neurons, and also showing how one such gene affects sleep and translation and a second affects sleep and memory.

Public Reviews:

Reviewer #2 (Public review):

Prior work by the Sehgal group has shown that a small group of neurons in the fly brain (anterior posterior (ap) α'β' mushroom body neurons (MBNs)) promote sleep and sleep-dependent appetitive memory specifically under fed conditions (Chouhan et al., (2021) Nature). Here, Li, Chouhan et al. combine cell-specific transcriptomics with measurements of sleep and memory to identify molecular processes underlying this phenomenon. They define transcriptional changes in ap α'β' MBNs and suggest a role for two genes downregulated following memory induction (Polr1F and Regnase-1) in regulating sleep and memory.

The transcriptional analyses in this manuscript are impressive. The authors have now included additional experiments that define acute and developmental roles for Polr1F and Regnase-1 respectively in regulating sleep. They have also provided additional data to strengthen their conclusion that Polr1F knockdown in α'β' mushroom body neurons enhances sleep.

The resubmitted work represents a convincing investigation of two novel sleep-regulatory proteins that may also play important roles in memory formation.

The authors have comprehensively addressed my comments, which I very much appreciate. I congratulate them on this excellent work.

We very much appreciate the reviewer’s positive feedback. Thank you!

Reviewer #3 (Public review):

Previous work (Chouhan et al., 2022) from the Sehgal group investigated the relationship between sleep and long-term memory formation by dissecting the role of mushroom body intrinsic neurons, extrinsic neurons, and output neurons during sleep-dependent and sleep-independent memory consolidation. In this manuscript, Li et al., profiled transcriptome in the anterior-posterior (ap) α'/β' neurons and identified genes that are differentially expressed after training in fed condition, which supports sleep-dependent memory formation. By knocking down candidate genes systematically, the authors identified Polr1F and Regnase-1 as two important hits that play potential roles in sleep and memory formation. What is the function of sleep and how to create a memory are two long-standing questions in science. The present study used a new approach to identify novel components that may link sleep and memory consolidation in a specific type of neuron. Importantly, these components implicated that RNA processing may play a role in these processes.

While I am enthusiastic about the innovative approach employed to identify RNA processing genes involved in sleep regulation and memory consolidation, I feel that the data presented in the manuscript is insufficient to support the claim that these two genes establish a definitive link between sleep and memory consolidation. First, the developmental role of Regnase-1 in reducing sleep remains unclear because knocking down Regnase-1 using the GeneSwitch system produced neither acute nor chronic sleep loss phenotype. In the revised manuscript, the author used the Gal80ts to restrict the knockdown of Regnase-1 in adult animals and concluded that Regnase-1 RNAi appears to affect sleep through development. Conducting overexpression experiments of Regnase-1 would lend some credibility to the phenotypes, however, this is not pursued in the revised manuscript. Second, while constitutive Regnase-1 knockdown produced robust phenotypes for both sleep-dependent and sleep-independent memory, it also led to a severe short-term memory phenotype. This raises the possibility that flies with constitutive Regnase-1 knockdown are poor learners, thereby having little memory to consolidate. The defect in learning could be simply caused by chronic sleep loss before training. Thus, this set of results does not substantiate a strong link between sleep and long-term memory consolidation. Lastly, the discussion on the sequential function of training, sleep, and RNA processing on memory consolidation appears speculative based on the present data.

We thank the reviewer for the enthusiasm about the approach. As noted above, we have now removed all claims about a link between sleep and memory, and instead just emphasize that we have identified RNA processing genes that affect sleep and memory. We agree with the reviewer that the basis of the Regnase-1 memory phenotype is unclear as the flies may be poor learners. Also, the learning/memory defects could be secondary to sleep loss or, as Reviewer 4 below suggests, all the behavioral deficits could be caused by impaired development/function of the relevant ap ɑ′/β′ cells. We have now included this possibility in the discussion of the manuscript. And we have modified the discussion on training, RNA processing, sleep and memory to emphasize the need for future experiments to address the sequence and relationship of these different processes.

Reviewer #4 (Public review):

Summary:

Li and Chouhan et al. follow up on a previous publication describing the role of anterior-posterior (ap) and medial (m) ɑ′/β′ Kenyon cells in mediating sleep-dependent and sleep-independent memory consolidation, respectively, based on feeding state in Drosophila melanogaster. The authors sequenced bulk RNA of ap ɑ′/β′ Kenyon cells 1h after flies were either trained-fed, trained-starved or untrained-fed and find a small number of genes (59) differentially expressed (3 upregulated, 56 downregulated) between trained-fed and trained-starved conditions. Many of these genes encode proteins involved in the regulation of gene expression. The authors then screened these differentially expressed genes for sleep phenotypes by expressing RNAi hairpins constitutively in ap ɑ′/β′ Kenyon cells and measuring sleep patterns. Two hits were selected for further analysis: Polr1F, which promoted sleep, and Regnase-1, which reduced sleep. The pan-neuronal expression of Polr1F and Regnase-1 RNAi constructs was then temporally restricted to adult flies using the GeneSwitch system. Polr1F sleep phenotypes were still observed, while Regnase-1 sleep phenotypes were not, indicating developmental defects. Appetitive memory was then assessed in flies with constitutive knockdown of Polr1F and Regnase-1 in ap ɑ′/β′ Kenyon cells. Polr1F knockdown did not affect sleep-dependent or sleep-independent memory, while Regnase-1 knockdown disrupted sleep-dependent memory, sleep-independent memory, as well as learning. Polr1F knockdown increased pre-ribosomal RNA transcripts in the brain, as measured by qPCR, in line with its predicted role as part of the RNA polymerase I complex. A puromycin incorporation assay to fluorescently label newly synthesized proteins also indicated higher levels of bulk translation upon Polr1F knockdown. Regnase-1 knockdown did not lead to observable changes in measurements of bulk translation.

Strengths:

The proposed involvement of RNA processing genes in regulating sleep and memory processes is interesting, and relatively unexplored. The methods are satisfactory.

Weaknesses:

The main weakness of the paper is in the overinterpretation of their results, particularly relating to the proposed link between sleep and memory consolidation, as stated in the title. Constitutive Polr1F knockdown in ap ɑ′/β′ Kenyon cells had no effect on appetitive long-term memory, while constitutive Regnase-1 knockdown affected both learning and memory. Since the effects of constitutive Regnase-1 knockdown on sleep could be attributed to developmental defects, it is quite plausible that these same developmental defects are what drive the observed learning and memory phenotypes. In this case, an alternative explanation of the authors' findings is that constitutive Regnase-1 knockdown disrupts the entire functioning of ap ɑ′/β′ Kenyon cells, and as a consequence behaviors involving these neurons (i.e. learning, memory and sleep) are disrupted. It will be important to provide further evidence of the function of RNA processing genes in memory in order to substantiate the memory link proposed by the authors.

As noted above, we have removed claims of a link between sleep and memory and instead focused the manuscript on our findings of RNA processing genes modulated during sleep-dependent memory. We concur that impaired development of ap ɑ′/β′neurons could account for the sleep and memory phenotype observed and have included this possibility in the manuscript.

Recommendations for the authors:

Reviewer #4 (Recommendations for the authors):

The title of the paper should be reconsidered to reflect the results. The evidence for a link between RNA processing genes and memory is weak.

We have changed the title.

Line 328. The term "central dogma" is misused. The central dogma refers to the unidirectional flow of information from DNA to protein. Instead the authors mean "gene expression".

Changed, thank you.

A couple of minor comments relating to the figures:

Figure 1b. It is not clear what the number 10570 in the bottom right corner refers to.

Fixed.

Figure 3b. RU- and RU+ annotation is missing (as shown in 3d).

Fixed.

Changes from the earlier version to the revised version

Abstract

Abstract

Earlier version: Memory consolidation in Drosophila can be sleep-dependent or sleep-independent, depending on the availability of food. Different regions of the mushroom body (MB) mediate these two mechanisms, with the anterior posterior (ap) alpha′/beta′ (α′/β′) neurons required for sleep-dependent memory consolidation in flies that are fed after training. These neurons are also involved in the increase of sleep after training, suggesting a coupling of sleep and memory. To better understand the mechanisms underlying sleep and memory consolidation initiation, we analyzed the transcriptome of ap α′/β′ neurons one hour after appetitive memory conditioning. A small number of genes were differentially expressed specifically in flies fed after training, but not in trained and starved flies or untrained flies. Knockdown of each of these differentially expressed genes in the ap α′/β′ neurons revealed multiple genes that affect sleep, with notable effects observed for Polr1F and Regnase-1, both of which decrease in expression after conditioning. Knockdown of Polr1F, a regulator of ribosome RNA transcription, in adult flies promotes sleep and increases pre-ribosome RNA expression as well as overall translation, supporting a function for Polr1F downregulation in memory consolidation. Conversely, knockdown of Regnase-1, an mRNA decay protein localized to the ribosome, reduces sleep. Given that Regnase-1 knockdown in ap α′/β′ neurons affects both sleep-dependent and sleep-independent memory, as well as short-term memory, Regnase-1 likely has an early role in the learning process, which may obscure a later function for its downregulation during sleep-dependent memory. These findings indicate that changes in RNA processing play a crucial role in triggering post-training sleep and memory consolidation.

Revised version: Memory consolidation in Drosophila can be sleep-dependent or sleep-independent, depending on the availability of food. The anterior posterior (ap) alpha′/beta′ (α′/β′) neurons of the mushroom body (MB) are required for sleep-dependent memory consolidation in flies fed after training. These neurons are also involved in the increase of sleep after training, suggesting a coupling of sleep and memory. To better understand the mechanisms underlying sleep and memory consolidation initiation, we analyzed the transcriptome of ap α′/β′ neurons one hour after appetitive memory conditioning. A small number of genes, enriched in RNA processing functions, were differentially expressed in flies fed after training relative to trained and starved flies or untrained flies. Knockdown of each of these differentially expressed genes in the ap α′/β′ neurons revealed notable sleep phenotypes for Polr1F and Regnase-1, both of which decrease in expression after conditioning. Knockdown of Polr1F, a regulator of ribosome RNA transcription, in adult flies promotes sleep and increases pre-ribosome RNA expression as well as overall translation, supporting a function for Polr1F downregulation in sleep-dependent memory. Conversely, while constitutive knockdown of Regnase-1, an mRNA decay protein localized to the ribosome, reduces sleep, adult specific knockdown suggests that effects of Regnase-1 on sleep are developmental in nature. We further tested the role of each gene in memory consolidation. Knockdown of Polr1F does not affect memory, which may be expected from its downregulation during memory consolidation. Regnase-1 knockdown in ap α′/β′ neurons impairs all memory, including short-term, implicating Regnase-1 in memory, but leaving open the question of why it is downregulated during sleep-dependent memory. Overall, our findings demonstrate that the expression of RNA processing genes is modulated during sleep-dependent memory and, in the case of Polr1F, this modulation likely contributes to increased sleep.

Background

Sleep is an optimized state for memory consolidation compared to wake (Rasch and Born, 2013). Indeed, sleep is thought to reorganize and strengthen the neural connections required for novel memory formation and long-term memory consolidation, and sleep disruption leads to impaired memory consolidation (Roselli et al., 2021). Beneficial effects of sleep on memory have also been attributed to post-learning neuronal reactivation (Wagner et al., 2007; Dag et al., 2019).

We recently demonstrated that Drosophila switch between sleep-dependent and sleep-independent memory consolidation based on food availability (Chouhan et al., 2021). Flies that are fed after appetitive conditioning show sleep-dependent memory consolidation, which is mediated by the anterior posterior (ap) α’/β’ neurons of the mushroom body. On the other hand, flies starved after training display sleep-independent memory mediated by medial (m) neurons of the α’/β’ lobes. Neurotransmission from ap α′/β′ neurons in the first 4 hours after training in the fed flies is not only required for long term memory but also increased sleep post-training. Indeed, activation of ap α′/β′ neurons promotes baseline sleep, supporting the idea that post training sleep is triggered by the same neurons that are required for sleep-dependent memory consolidation. However, how sleep and memory consolidation are coordinated in ap α′/β′ neurons in this time window is not understood.

Earlier version: Since both post-training sleep and memory consolidation occur rapidly within a few hours after training, we investigated transient gene expression changes in the ap α′/β′ neurons after training to elucidate the interplay of sleep and memory consolidation. Here, we profiled the transcriptome of ap α′/β′ neurons 1 hour after flies were fed and trained, as well as under control conditions in which flies were starved and trained or fed and untrained. We knocked down the differentially expressed genes in trained and fed flies and identified two RNA processing genes that affect sleep. Knockdown of one of these, Polr1F, a regulator of ribosomal RNA synthesis, promotes sleep and translation, both of which are required for consolidation of memory in trained and fed flies. Knockdown of Polr1F does not affect memory, which is consistent with downregulation of this gene during memory consolidation. Knockdown of Regnase-1, an mRNA decay protein, reduces sleep and disrupts short- and long-term memory, suggesting an early role, perhaps in learning. These results suggest that RNA processing is an important mechanism linking sleep to memory.

Revised version: Since both post-training sleep and memory consolidation occur rapidly within a few hours after training, we investigated transient gene expression changes in the ap α′/β′ neurons after training to elucidate the interplay of sleep and memory consolidation. Here, we profiled the transcriptome of ap α′/β′ neurons 1 hour after flies were fed and trained, as well as under control conditions in which flies were starved and trained or fed and untrained. We knocked down the differentially expressed genes in trained and fed flies and identified two RNA processing genes that affect sleep. Knockdown of one of these, Polr1F, a regulator of ribosomal RNA synthesis, promotes sleep and translation, both of which are required for consolidation of memory in trained and fed flies. Knockdown of Polr1F does not affect memory, which is consistent with downregulation of this gene during memory consolidation. Knockdown of Regnase-1, an mRNA decay protein, reduces sleep and disrupts post-training sleep, and also impairs both short- and long-term memory. While the relevance of Regnase-1 downregulation during sleep-dependent memory is unclear, overall we propose that RNA processing is modulated during sleep-dependent memory, likely to regulate sleep and memory.

Results

Transient transcriptome profiling of ap α′/β′ neurons after training

We previously demonstrated that flies fed after appetitive memory training exhibit increased sleep and form sleep-dependent memory, which is mediated by ap α′/β′ neurons of the mushroom body (MB) α’/β’ lobes (Chouhan et al., 2021). To address the mechanisms that mediate increases in sleep and consolidation of memory in ap α′/β′ neurons, we assayed gene expression changes in ap α′/β′ neurons of flies fed after training in an appetitive conditioning paradigm. We crossed ap α′/β′ neuron driver R35B12-Gal4 (BDSC #49822) with UAS-nGFP (BDSC #4775) to label all ap α′/β′ neurons with nuclear GFP and collected 5-7 day old F1 progeny, which were subjected to one of three different conditions: Trained-Fed, Trained-Starved, Untrained-Fed as illustrated (Figure 1A). The trained-starved flies served as controls for sleep-dependent changes, while flies that were fed but untrained served as controls for training-dependent changes. After one-hour, 50 mixed sex (25 for each sex) fly brains from each condition were dissected, and 500 GFP+ cells were sorted for bulk-RNA sequencing using the protocol described by Hongjie Li et al (Li et al., 2017).

Differential gene expression after training in mushroom ap α′/β′ neurons.

(A) 5-7 day old mixed sex wCS flies were exposed to one of the following three conditions: Trained-Fed, Trained-Starved and Untrained-Fed. Only Trained-Fed flies are expected to increase sleep after treatment and thus form sleep-dependent memory (Chouhan et al., 2021). Brain dissection, single cell suspension and cell sorting were used to extract ap α′/β′ neurons in each of these three different conditions, and bulk-sequencing of the sorted cells was conducted. (B) We sequenced four samples for each condition and subjected them to differential gene analysis of pairwise comparison of the three conditions. We found that 59 genes are significantly different in TrainedStarved vs. TrainedFed and UntrainedFed vs. TrainedFed. (C) Gene ontology (GO) analysis of these 59 genes using FlyEnrichr (https://maayanlab.cloud/FlyEnrichr/) revealed that they encode cellular components of the 90S preribosome, Cajal body, DNA-directed RNA polymerase complex, nuclear euchromatin, and condensed chromosome.(D) Heatmap of the 59 DEGs including two genes labeled in red, CG13773 (Polr1F) and Regnase-1, which affect sleep and are the focus of this study due to their impacts on sleep. DEGs were identified by DESeq2 with the cutoff of FDR< 0.1 and fold change >1.5. Heatmaps were plotted by using TPM values of genes for each sample; data were log-transformed and scaled row-wise for visualization.

Our analysis of the bulk RNA-seq data revealed that most genes are not altered in ap α′/β′ neurons after one hour, so there are no significant global changes observed among the three conditions. The correlation matrix calculated using the top 75% genes showed high similarity between samples, with most values exceeding 0.9 (Figure 1 – source data). Nonetheless, we did observe a small subset of genes that were rapidly responsive and differentially expressed between the groups (Figure 1B). Pathway Principal Component Analysis (PathwayPCA) of these differentially expressed genes (DEGs), a statistical method to identify key patterns in gene expression data by reducing data dimensionality, showed that genes contributing to PC1, which accounted for 35% of the variance, largely encode proteins involved in transcription and biosynthesis, including RNA biosynthesis processes (Figure 1 – supplemental figure 1A) (Odom et al., 2020). This suggests that the training and feeding paradigm influenced transcription and RNA biosynthetic processes. Meanwhile, gene ontology analysis of these DEGs, performed using DAVID Bioinformatics (https://david.ncifcrf.gov/home.jsp), revealed significant enrichment in processes such as protein unfolding/refolding and the stress response (Figure 1 – supplemental figure 1B) (Sherman et al., 2022).

Of all the DEGs across the three groups, we focused on 59 DEGs that were significantly different in TrainedStarved vs. TrainedFed and UntrainedFed vs. TrainedFed; of these, 56 were downregulated and only three were upregulated in the Trained-Fed condition compared to the control conditions (Figure 1B, D, Figure 1 – source data). Gene ontology (GO) analysis of these 59 genes using FlyEnrichr (https://maayanlab.cloud/FlyEnrichr/) indicated that they encode cellular components of the 90S preribosome, Cajal body, DNA-directed RNA polymerase complex, nuclear euchromatin, and condensed chromosome, consistent with the PCA enrichment (Figure 1C) (Chen et al., 2013). Our transcriptome results suggest that ribosome biosynthesis and transcription are the initial changes in ap α′/β′ neurons of trained and fed flies.

Two genes expressed differentially in ap α′/β′ neurons of trained and fed flies predominately affect sleep

Earlier version: To investigate if any of the 59 DEGs identified in the ap α′/β′ neurons of trained and fed flies affect baseline sleep, we knocked down each of the genes using UAS-RNAi lines and screened for their potential effects on sleep. We used the ap α′/β′ neuron constitutive driver R35B12-Gal4 line to drive the RNAi constructs and compared the sleep patterns of knockdown flies with those of R35B12-Gal4 and UAS-RNAi control flies (Figure 2A, B, Figure 2 – source data). Using a cutoff of a 200 min change in sleep relative to each control group, we identified two genes, Polr1F and Regnase-1, that showed significant effects on baseline sleep. Knockdown of Polr1F, a component of RNA polymerase I complex, led to an increase in sleep. Although Polr1F has not been extensively studied in Drosophila (Marygold et al., 2020), its human ortholog hRPA43 is part of the multi-subunit protein complex Pol I that regulates the transcription of ribosomal RNA (Beckouët et al., 2011). Knockdown of Regnase-1, an RNA-binding protein that binds to mRNA undergoing active translation and promotes mRNA decay via its ribonuclease activity, in ap α′/β′ neurons resulted in a reduction of both nighttime and daytime sleep (Figure 2C-F). Further analysis of sleep architecture revealed that knockdown of Polr1F and Regnase-1 did not significantly impact the total activity of flies, while knockdown of either Polr1F and Regnase-1 resulted in increased and decreased average length of sleep episodes, respectively (Figure 2 – supplemental figure1A-F).

Revised version: To investigate if any of the 59 DEGs identified in the ap α′/β′ neurons of trained and fed flies affect baseline sleep, we knocked down each of the genes using UAS-RNAi lines and screened for their potential effects on sleep. We used the ap α′/β′ neuron constitutive driver R35B12-Gal4 line to drive the RNAi constructs and compared the sleep patterns of knockdown flies with those of R35B12-Gal4 and UAS-RNAi control flies (Figure 2A, B, Figure 2 – source data). Using a cutoff of a 200 min change in sleep relative to each control group, we identified two genes, Polr1F and Regnase-1, that showed significant effects on baseline sleep. Knockdown of Polr1F, a component of RNA polymerase I complex, led to an increase in sleep. Although Polr1F has not been extensively studied in Drosophila (Marygold et al., 2020), its human ortholog hRPA43 is part of the multi-subunit protein complex Pol I that regulates the transcription of ribosomal RNA (Beckouët et al., 2011). Knockdown of Regnase-1, an RNA-binding protein that binds to mRNA undergoing active translation and promotes mRNA decay via its ribonuclease activity, in ap α′/β′ neurons resulted in a reduction of both nighttime and daytime sleep (Figure 2C-F). Further analysis of sleep architecture revealed that knockdown of Polr1F and Regnase-1 did not significantly impact the total activity of flies, while knockdown of either Polr1F and Regnase-1 resulted in increased and decreased average length of sleep episodes, respectively (Figure 2 – supplemental figure 1A-F).

The sleep screen of differentially expressed genes identifies Polr1F and Regnase-1 as sleep-regulating genes.

(A-B) Flies carrying the ap α′/β′ neuron driver R35B12-Gal4 were crossed with flies carrying UAS-RNAi constructs targeting DEGs identified from RNA-seq analysis. 5-7 days old female F1 progeny were loaded onto Trikinetics DAM monitors to measure their sleep in a 12-hour light: 12-hour dark (12:12 LD) cycle. Mean total sleep (A) and nighttime sleep (B) were calculated by Pysolo and the difference between experimental flies and Gal4 and RNAi controls was calculated separately for each independent experiment; average values comparing each experimental to its Gal4 control (X-axis) and RNAi control (Y-axis) are shown in the plots. Of all the lines screened, knockdown of Polr1F and Regnase-1 had strongest effects on sleep, producing an increase and decrease in sleep respectively. (C-F) show the representative sleep traces of R35B12-Gal4>polr1FRNAi flies and R35B12-Gal4>regnase1RNAi. N = 23-32 per genotype from two independent replicates combined are shown in E and F respectively, and bar graphs show mean + SEM. p values for each comparison were calculated using the Kruskal-Wallis test with Dunn’s multiple comparisons test. **p<0.01, ***p<0.001, ****p<0.0001.

Earlier version: Given that these two genes reduce expression rapidly in response to training under fed conditions, we next used the inducible pan-neuronal GeneSwitch driver nSyb-GeneSwitch (GS) to determine if restricting knockdown of these two genes to the adult stage recapitulates changes in sleep and/or memory. RU486 (mifepristone) was added to normal fly food and transgene expression is induced when flies are loaded into Drosophila Activity Monitor (DAM) glass tubes (Robles-Murguia et al., 2019). We crossed the nSyb-GS flies with Polr1F or Regnase-1 RNAi flies and transferred the adult F1 progeny to RU486 tubes to induce expression of the RNAi 3-5 hours before dusk, and continuously monitored sleep for 5 days. Knockdown of Polr1F resulted in immediate inactivity and sleep (Figure 3A-C). However, with knockdown of Regnase-1, immediate changes in sleep were not noted (Figure 3 – supplemental figure 3A-C). This could be due to potential leakiness of the nSyb-GS, such that it reduced sleep even without RU486 treatment, coupled with the fact that sleep before dusk is already quite low, potentially hindering detection of further sleep reduction. We also analyzed sleep for the next 3 consecutive days from day 3 to day 5 and found that constitutive pan-neuronal Polr1F knockdown increased sleep while Regnase-1 knockdown seemed to have no effect, again perhaps due to the leakiness of nSyb-GS. However, it is also possible that Regnase-1 affects sleep during developmental stages or acts differently in different brain areas (Figure 3D, E; Figure 3 – supplemental figure 3D, E). In general, these results indicate that adult specific knockdown of Polr1F promotes sleep while brain-wide adult-specific knockdown of Regnase-1 has limited effect on sleep.

Revised version: Given that these two genes reduce expression rapidly in response to training under fed conditions, we next used the inducible pan-neuronal GeneSwitch driver nSyb-GeneSwitch (GS) to determine if restricting knockdown of these two genes to the adult stage recapitulates changes in sleep and/or memory. RU486 (mifepristone) was added to normal fly food and transgene expression is induced when flies are loaded into Drosophila Activity Monitor (DAM) glass tubes (Robles-Murguia et al., 2019). We crossed the nSyb-GS flies with Polr1F or Regnase-1 RNAi flies and transferred the adult F1 progeny to RU486 tubes to induce expression of the RNAi 3-5 hours before dusk, and continuously monitored sleep for 5 days. Knockdown of Polr1F resulted in immediate inactivity and sleep (Figure 3A-C). However, with knockdown of Regnase-1, immediate changes in sleep were not noted (Figure 3 – supplemental figure 2A-C).

Acute and chronic effects of pan-neuronal knockdown of Polr1F on sleep in adult flies.

(A) Schematic representation of transient and chronic sleep measurements in nSyb-GS>polr1fRNAi flies. (B) Representative sleep traces and transient activity plot of flies expressing Polr1F RNAi under the control of an inducible pan-neuronal driver (nSyb-GS>polr1f RNAi) with and without RU treatment. (C) Quantification of sleep during the first 3 hours (ZT 9-12) after F1 progeny flies were loaded into RU- or RU+ DAM tubes at ZT8-T9. Sleep was measured starting at ZT9. N = 39-40 individual flies per replicate with data from three independent replicates combined. The Mann-Whitney test was used to compare RU+ group and RU-groups. (D) Representative average sleep traces of nSyb-GS>polr1f RNAi in the RU- and RU+ DAM tubes for three consecutive days. Chronic sleep effects of pan-neuronal knockdown Polr1F were measured based on sleep data from day 3 to day 5. (E) Quantification of average total sleep of nSyb-GS>polr1fRNAi and controls in the DAM tubes from (D). Unpaired t-test was used to compare between RU- and RU+ groups. ****p<0.0001.

Added in the revised version: This could be due to potential leakiness of the nSyb-GS, such that it reduced sleep even without RU486 treatment, coupled with the fact that sleep before dusk is already quite low, potentially hindering detection of further sleep reduction. We also analyzed sleep for the next 3 consecutive days from day 3 to day 5 and found that constitutive pan-neuronal Polr1F knockdown increased sleep while Regnase-1 knockdown seemed to have no effect, again perhaps due to the leakiness of nSyb-GS. However, it is also possible that Regnase-1 affects sleep during developmental stages or acts differently in different brain areas (Figure 3D, E; Figure 3 – supplemental figure 2D, E). In general, these results indicate that adult specific knockdown of Polr1F promotes sleep while brain-wide adult-specific knockdown of Regnase-1 has limited effect on sleep.

Earlier version: To further assess the adult-specific sleep function of both genes in ap α′/β′ neurons, we coupled the R35B12-Gal4 driver to tubulin-Gal80ts, which allows temperature-dependent expression of Gal4, and crossed it with the UAS-Polr1FRNAi and UAS-Regnase-1RNAi lines. We confirmed that Polr1F RNAi promotes both transient and chronic sleep when flies were transferred from permissive to restrictive temperatures during the adult stages (Figure 3 – supplemental figure 3). Conversely, Regnase-1 RNAi showed no effect on sleep in the adult stage, consistent with our nSyb-GS experiments, suggesting that the Regnase-1 RNAi sleep effect is likley developmental (Figure 3 – supplemental figure 3). We also analyzed the daily peak activity of the flies to determine if either Polr1F RNAi or Regnase-1 RNAi affected activity. We found that neither gene affected peak activity, suggesting that both genes influence sleep but not activity (Figure 3 - supplemental figure 4).

Revised version: To further assess the adult-specific sleep function of both genes in ap α′/β′ neurons, we coupled the R35B12-Gal4 driver to tubulin-Gal80ts, which allows temperature-dependent expression of Gal4, and crossed it with the UAS-Polr1FRNAi and UAS-Regnase-1RNAi lines. We confirmed that Polr1F RNAi promotes both transient and chronic sleep when flies were transferred from permissive to restrictive temperatures during the adult stages (Figure 3 – supplemental figure 1). Conversely, Regnase-1 RNAi showed no effect on sleep in the adult stage, consistent with our nSyb-GS experiments, suggesting that the Regnase-1 RNAi sleep effect is likely developmental (Figure 3 – supplemental figure 3). We also analyzed the daily peak activity of the flies to determine if either Polr1F RNAi or Regnase-1 RNAi affected activity. We found that neither gene affected peak activity, suggesting that both genes influence sleep but not activity (Figure 3 -supplemental figure 4).

Earlier version: Knockdown of Regnase-1 affects memory

Revised version: Knockdown of Regnase-1 affects memory and post-training sleep

Earlier version: We next evaluated the impact of Polr1F and Regnase-1 knockdown on memory consolidation using our olfactory conditioning paradigm (Figure 4A). Starved flies were subjected to training to associate an odor with a reward, and then post-training, they were either kept on food vials for sleep-dependent memory consolidation or kept starved to promote sleep-independent memory consolidation. Memory tests were conducted 24 hours after training for starved flies, while fed flies were restarved for 42 hours before testing, as starvation is necessary for memory retrieval (Krashes and Waddell, 2008). Efficacy of the RNAi line used above was verified by qPCR experiments (Figure 4 -supplemental figure 1A). We observed that constitutive knockdown of Polr1F in ap α′/β′ neurons did not affect sleep-dependent or sleep-independent memory as memory performance was comparable to that of genetic controls (Figure 4B-C). These results were consistent with the fact that Polr1F levels typically decrease during memory consolidation. Monitoring of sleep from ZT8 to ZT12 after training at zeitgeber time (ZT) 6 showed that the post-training increase in sleep was also not affected by Polr1F knockdown in ap α′/β′ neurons (Figure 4E), suggesting that acute downregulation of Polr1F may not be essential to the post-training increased sleep.

Revised version: We next evaluated the impact of Polr1F and Regnase-1 knockdown on memory consolidation using our olfactory conditioning paradigm (Figure 4A). Starved flies were subjected to training to associate an odor with a reward, and then post-training, they were either kept on food vials for sleep-dependent memory consolidation or kept starved to promote sleep-independent memory consolidation. Memory tests were conducted 24 hours after training for starved flies, while fed flies were restarved for 42 hours before testing, as starvation is necessary for memory retrieval (Krashes and Waddell, 2008). Efficacy of the RNAi line used above was verified by qPCR experiments (Figure 4 -supplemental figure 1A). We observed that constitutive knockdown of Polr1F in ap α′/β′ neurons did not affect sleep-dependent or sleep-independent memory as memory performance was comparable to that of genetic controls (Figure 4B-C). These results were consistent with the fact that Polr1F levels typically decrease during memory consolidation. Monitoring of sleep from ZT8 to ZT12 after training at zeitgeber time (ZT) 6 showed that the post-training increase in sleep was also not affected by Polr1F knockdown in ap α′/β′ neurons (Figure 4E), suggesting that Polr1F does not have to be acutely downregulated for post-training sleep.

Regnase-1 expression is essential for sleep-dependent and sleep-independent memory.

(A) Schematic representation of the memory test protocol. (B, C) Sleep-dependent and sleep-independent memory tests were conducted under fed and starved conditions, respectively. Knockdown of Regnase-1 significantly reduces long-term memory performance in both fed and starved flies. However, knockdown of Polr1F in ap α′/β′ neurons does not affect long-term memory performance. N≥6 biological replicates, each replicate containing 100-150 flies. (D, E) Fed UAS-regnase-1-RNAi/+ and R35B12/+ flies exhibit a significant increase in sleep after training, while R35B12-Gal4>regnase-1 RNAi flies fail to show a comparable increase in post-training sleep. The total sleep in the ZT8-ZT12 interval is shown in (E). Polr1F knockdown in ap α′/β′ neurons does not affect the post-training increase in sleep. N≥32. (F) Compared to R35B12-Gal4/+ and +/UAS-Regnase-1 RNAi flies, R35B12-Gal4>reganse-1 RNAi flies show a significant decrease in the performance index in short-term memory. N≥6 biological replicates, each containing 100-150 flies. ns=not significant, p>0.05, *p<0.05, **p<0.01, ***p<0.001, ****p<0.0001.

Earlier version: On the other hand, constitutive Regnase-1 knockdown in ap α′/β′ neurons resulted in a significant decrease in long-term memory performance in both fed and starved flies and eliminated the increase in post-training sleep (Figure 4B-E). These findings suggested that Regnase-1 expression in ap α′/β′ neurons is necessary for both sleep-dependent and sleep-independent memory consolidation. It is also possible that disruption of Regnase-1 in ap α′/β′ neurons affects learning, short-term memory formation, or appetitive memory retrieval (Aso et al., 2014; Shyu et al., 2019; Rutherford et al., 2023). Indeed, short-term memory tests confirmed that Regnase-1 knockdown flies performed significantly worse than control flies (Figure 4F), indicating that Regnase-1 acts early in the learning/memory process, which might explain its effect on sleep-dependent and sleep-independent memory. We also tested Regnase-1 overexpression for effects on sleep, and found that it does not affect sleep (Figure 4 - supplemental figure 1B) (Zhu et al., 2019).

Revised version: On the other hand, constitutive Regnase-1 knockdown in ap α′/β′ neurons resulted in a significant decrease in long-term memory performance in both fed and starved flies and eliminated the increase in post-training sleep (Figure 4B-E). These findings suggested that Regnase-1 expression in ap α′/β′ neurons is necessary for both sleep-dependent and sleep-independent memory consolidation. It is also possible that disruption of Regnase-1 in ap α′/β′ neurons affects learning, short-term memory formation, or appetitive memory retrieval (Aso et al., 2014; Shyu et al., 2019; Rutherford et al., 2023). Indeed, short-term memory tests confirmed that Regnase-1 knockdown flies performed significantly worse than control flies (Figure 4F). Given that the sleep effect of Regnase-1 is developmental, it is possible that this is also the case for the memory phenotype.

Added in the revised version: Regnase-1 overexpression had no effect on sleep (Figure 4 -supplemental figure 1B) (Zhu et al., 2019), but knockdown by R26E01-Gal4 in the m α′/β′ neurons also reduced sleep (data not shown), suggesting that Regnase-1 has a broader role that is not restricted to sleep-dependent memory.

Knockdown of Polr1F promotes translation

Since Polr1F knockdown promotes sleep and a 22 amino acid peptide within Polr1F inhibits ribosomal DNA transcription (Rothblum et al., 2014), we predicted that Polr1F acted as a suppressor for ribosomal DNA transcription, and thus knockdown of Polr1F would enhance the transcription and translation of ribosomal RNA and thereby overall protein synthesis. As Regnase-1 is thought to promote decay of mRNAs undergoing translation, its knockdown, or its downregulation following training, might also be expected to promote translation, or at least the translation of its target mRNAs. The role of protein synthesis in long-term memory consolidation is well-established across organisms (Alberini and Kandel, 2015), and nascent rRNA synthesis was also shown to be induced by training and required for memory consolidation in mouse (Allen et al., 2018). We thus used real-time qPCR of total RNA extracted from vehicle control and RU-treated fly brains to measure how precursor ribosomal RNA (Pre-rRNA) is affected by knockdown of Polr1F. We found that the pre-rRNA level increased significantly in the RU486 induction (RU+) group compared with vehicle control (RU-) group for nSyb-GS>polr1F RNAi flies (Figure 5A), indicating Polr1F knockdown results in higher levels of pre-rRNA, which is consistent with studies of Polr1F homologs in yeast cells (Thuriaux et al., 1995; Rothblum et al., 2014).

Knockdown of Polr1F results in high translation.

Earlier version: (A) The nSyb-GS>polr1f RNAi flies exhibit a significant increase in pre-rRNA levels. (B) The ex-vivo puromycin immunostaining assay was used to measure translation in dissected whole brains. The results show that knockdown of Polr1F using the pan-neuronal nSyb-GeneSwitch (GS) system increases translation relative to control flies that were not treated with RU. The normalized mean grayscales from the RU- and RU+ groups are compared using an unpaired t-test. The analysis includes data from 8-11 flies per group, with results from two independent replicates combined. Scale bar: 100 µm. (C) The schematic model illustrates the roles of Polr1F and Regnase-1 in memory consolidation. The genes Polr1F and Regnase-1 are prominently downregulated during memory consolidation in trained and fed flies, respectively. Polr1F is involved in regulating ribosomal RNA synthesis, and its decrease in levels in trained and fed flies promotes sleep and translation. In contrast, Regnase-1 is involved in mRNA decay, and its downregulation during memory consolidation may contribute to the stabilization of mRNAs that encode proteins important for long-term memory formation.

Revised version: (A) The nSyb-GS>polr1f RNAi flies exhibit a significant increase in pre-rRNA levels. (B) The ex-vivo puromycin immunostaining assay was used to measure translation in dissected whole brains. The results show that knockdown of Polr1F using the pan-neuronal nSyb-GeneSwitch (GS) system increases translation relative to control flies that were not treated with RU. The normalized mean grayscales from the RU- and RU+ groups are compared using an unpaired t-test. The analysis includes data from 8-11 flies per group, with results from two independent replicates combined. Scale bar: 100 µm. (C) The schematic model illustrates the major transcriptome features of ap α′/β′ neurons under trained and fed conditions, revealed in this manuscript. Two RNA processing genes Polr1F and Regnase-1 are prominently downregulated during memory consolidation in trained and fed flies, respectively. Polr1F is involved in regulating ribosomal RNA synthesis, and its decrease in levels in trained and fed flies promotes sleep and translation. In contrast, Regnase-1 is localized in the ribosome and involved in mRNA decay, and its downregulation causes deficits in sleep and memory.

Earlier version: The increasing ribosomal RNA should help translation and so we also used incorporation of puromycin into newly synthesized peptides as an estimate of translation after inducing pan-neuronal knockdown of Polr1F. We observed significantly higher levels of anti-puromycin GFP fluorescence when Porl1F was knockdown pan-neuronally (RU+) compared to the control (RU-) group (Figure 5B), indicating that Polr1F suppresses translation, probably by suppressing ribosomal RNA transcription (Rothblum et al., 2014). Thus, it may need to be downregulated after training to support translation and memory. Given that knockdown of Porl1F enhances rRNA synthesis and sleep, there is a question as to whether global alterations in rRNA synthesis impact sleep. However, feeding of an rRNA inhibitor (CX-5461) failed to produce rapid sleep phenotypes in fed or starved flies (Figure 5 – supplemental figure 5), suggesting that rRNA synthesis may not directly influence sleep. Using puromycin to address effects of Regnase-1 knockdown on translation revealed an insignificant slight increase in translation (data not shown); given that Regnase-1 may specifically affect pre-existing translationally active mRNA or may act only on specific target mRNAs, its effects on de novo translation may not be obvious.

Revised version: The increasing ribosomal RNA should help translation and so we also used incorporation of puromycin into newly synthesized peptides as an estimate of translation after inducing pan-neuronal knockdown of Polr1F. We observed significantly higher levels of anti-puromycin GFP fluorescence when Porl1F was knockdown pan-neuronally (RU+) compared to the control (RU-) group (Figure 5B), indicating that Polr1F suppresses translation, probably by suppressing ribosomal RNA transcription (Rothblum et al., 2014). Thus, it may need to be downregulated after training to support translation and memory. Given that knockdown of Porl1F enhances rRNA synthesis and sleep, there is a question as to whether global alterations in rRNA synthesis impact sleep. Feeding of an rRNA inhibitor (CX-5461) failed to produce rapid sleep phenotypes in fed or starved flies (Figure 5 – supplemental figure 1), but this negative result cannot be considered conclusive.

Added in the revised version: Using puromycin to address effects of Regnase-1 knockdown on translation revealed an insignificant slight increase in translation (data not shown); given that Regnase-1 may specifically affect pre-existing translationally active mRNA or may act only on specific target mRNAs, its effects on de novo translation may not be obvious.

Discussion

Earlier version: The anterior-posterior (ap) α′/β′ neurons of the mushroom body make critical and privileged contributions to sleep-dependent memory consolidation and post-training sleep (Krashes et al., 2007; Chouhan et al., 2021), but the mechanisms that link memory and sleep in these neurons are not known. To address this gap, we conducted transcriptomic analysis of ap α′/β′ neurons from trained and fed flies to identify genes that change rapidly under conditions that drive sleep-dependent memory. By uncovering two RNA processing genes involved in memory and sleep, our transcriptome profiling of ap α′/β′ neurons suggests that genes regulating rRNA transcription and translation are altered in the context of sleep-dependent memory consolidation.

Revised version: The anterior-posterior (ap) α′/β′ neurons of the mushroom body make critical and privileged contributions to sleep-dependent memory consolidation and post-training sleep (Krashes et al., 2007; Chouhan et al., 2021), but the mechanisms for sleep-dependent memory in these neurons are not known. To address this gap, we conducted transcriptomic analysis of ap α′/β′ neurons from trained and fed flies to identify genes that change rapidly under conditions that drive sleep-dependent memory. Our transcriptome profiling of ap α′/β′ neurons suggests that genes regulating rRNA transcription and translation are altered in the context of sleep-dependent memory consolidation.

RNA processing genes mediate sleep-dependent memory consolidation and sleep

Earlier version: Many of the 59 DEGs we identified are implicated in RNA processing. Of these, Polr1F and CG11920 affect ribosomal RNA processing, CG5654 is predicted to be part of the 90S pre-ribosome and involved in endonucleoytic cleavages (Herold et al., 2009), WDR79 encodes a small Cajal body specific RNA binding protein, Nup133 encodes a component of nuclear pore complex, Regnase-1 degrades mRNA, and CG18011, CG17568, Koko, CG11398 and Meics are all involved in RNA polymerase II-specific transcription (Di Giorgio et al., 2017; Rogg et al., 2022). The enrichment of RNA processing and translation genes is consistent with the notion that memory consolidation requires transcription and translation (Seibt and Frank, 2012; Alberini and Kandel, 2015). More importantly, our findings here indicate that alterations of some of RNA processing genes may impact sleep as well.

Revised version: Many of the 59 DEGs we identified are implicated in RNA processing. Of these, Polr1F and CG11920 affect ribosomal RNA processing, CG5654 is predicted to be part of the 90S pre-ribosome and involved in endonucleoytic cleavages (Herold et al., 2009), WDR79 encodes a small Cajal body specific RNA binding protein, Nup133 encodes a component of nuclear pore complex, Regnase-1 degrades mRNA, and CG18011, CG17568, Koko, CG11398 and Meics are all involved in RNA polymerase II-specific transcription (Di Giorgio et al., 2017; Rogg et al., 2022). The enrichment of RNA processing and translation genes is consistent with the notion that memory consolidation requires transcription and translation (Seibt and Frank, 2012; Alberini and Kandel, 2015). More importantly, our findings here indicate that alterations of some of RNA processing genes impact sleep as well.

Polr1F regulates ribosome RNA synthesis and sleep

Earlier version: Learning-induced changes in gene expression in memory-related neurons are often critical for long-term memory formation (Cavallaro et al., 2002; Hoedjes et al., 2015; Tadi et al., 2015). Our findings with Polr1F implicate changes in Pol I transcription during sleep-dependent memory. Polr1F(Rpa43) is predicted to be part of the RNA polymerase I complex and is involved in DNA-dependent RNA polymerase activity/rDNA transcription in yeast, especially for the initiation of ribosomal RNA (Beckouët et al., 2011; Marygold et al., 2020). As suggested by our data, Polr1F has an inhibitory role in the RNA polymerase I complex, which is also consistent with the aforementioned study that found a 22 amino acid peptide within Polr1F can inhibit rDNA transcription (Rothblum et al., 2014). Based upon our finding that inducible knockdown of Polr1F rapidly promotes translation, it is likely that the rapid and dramatic decline of Polr1F after training in fed flies serves to increase de novo ribosome RNA synthesis. This supports the report that ribosomal RNA is induced by learning and required for memory consolidation in mice (Allen et al., 2018). While we do not know how knockdown of Polr1F promotes sleep, an attractive possibility is that higher translation is a result of elevated sleep. Sleep is thought to promote translation and its role in memory consolidation in fed flies could be related to its effect on translation (Seibt and Frank, 2012; Chouhan et al., 2021). Alternatively, increased translation or rRNA synthesis could promote sleep. However, translation is typically thought of as a consequence of sleep rather than a cause (Zimmerman et al., 2006). Also, rRNA transcription rates remain constant throughout the day in the liver (Sinturel et al., 2017), although it is still possible that these rates vary in particular regions of the brain and affect sleep. The role of rRNA synthesis in Drosophila learning and memory has barely been explored, but our work, together with that of Allen et al. (2018), indicates that the well-known requirement for de novo protein synthesis during long-term memory consolidation (Jarome and Helmstetter, 2014) includes increased synthesis of ribosomal RNA and protein.

Revised version: Learning-induced changes in gene expression in memory-related neurons are often critical for long-term memory formation (Cavallaro et al., 2002; Hoedjes et al., 2015; Tadi et al., 2015). Our findings with Polr1F implicate changes in Pol I transcription during sleep-dependent memory. Polr1F(Rpa43) is predicted to be part of the RNA polymerase I complex and is involved in DNA-dependent RNA polymerase activity/rDNA transcription in yeast, especially for the initiation of ribosomal RNA (Beckouët et al., 2011; Marygold et al., 2020). As suggested by our data, Polr1F has an inhibitory role in the RNA polymerase I complex, which is also consistent with the aforementioned study that found a 22 amino acid peptide within Polr1F can inhibit rDNA transcription (Rothblum et al., 2014). Based upon our finding that inducible knockdown of Polr1F rapidly promotes translation, it is likely that the rapid and dramatic decline of Polr1F after training in fed flies serves to increase de novo ribosome RNA synthesis.

Added in the revised version: This supports the report that ribosomal RNA is induced by learning and required for memory consolidation in mice (Allen et al., 2018). While we do not know how knockdown of Polr1F promotes sleep, an attractive possibility is that higher translation is a result of elevated sleep. Sleep is thought to promote translation and its role in memory consolidation in fed flies could be related to its effect on translation (Seibt and Frank, 2012; Chouhan et al., 2021). Alternatively, increased translation or rRNA synthesis could promote sleep. However, translation is typically thought of as a consequence of sleep rather than a cause (Zimmerman et al., 2006). Also, rRNA transcription rates remain constant throughout the day in the liver (Sinturel et al., 2017), but it is still possible that these rates vary in particular regions of the brain and affect sleep. The role of rRNA synthesis in Drosophila learning and memory has barely been explored, but our work, together with that of Allen et al. (2018), indicates that the well-known requirement for de novo protein synthesis during long-term memory consolidation (Jarome and Helmstetter, 2014) includes increased synthesis of ribosomal RNA and protein.

Added in the revised version: Novel neuronal role of Regnase-1 in sleep and memory

Removed in the revised version: Regnase-1 inactivation may affect memory through effects on translation

Earlier version: The role of RNA binding protein Regnase-1 in the innate immune response has been extensively studied (Mino et al., 2015; Mao et al., 2017; Wei et al., 2019). However, our study sheds light on a novel function of Regase-1 in ap α′/β′ neurons on sleep and memory. Regnase-1 is an anti-inflammatory enzyme that inhibits mRNA translation during acute inflammatory responses. It localizes to the ribosomes on the surface of the endoplasmic reticulum (ER) and binds to translationally active mRNAs with specialized stem-loop structures at the 3′UTR (Uehata et al., 2013; Mino et al., 2015). When phosphorylated, Regnase-1 is released from the ER (Tanaka et al., 2019). Because functional Regnase-1 binds and degrades its bound mRNA, Regnase-1 inactivation leads to an increase of its target mRNA (Uehata et al., 2013). The target mRNAs of Regnase-1 in immune cells encode proinflammatory cytokines, which can then be expressed when Regnase-1 is inactivated. However, Regnase-1 has also been reported to modulate cytokines and neuronal injury in the microglia in rats (Liu et al., 2016). Regulation of Regnase-1 is usually rapid and transient, and its rapid response to microenvironmental changes, different pathological states and stress is critical for cellular adaption (Mao et al., 2017).

Revised version: The role of RNA binding protein Regnase-1 in the innate immune response has been extensively studied (Mino et al., 2015; Mao et al., 2017; Wei et al., 2019). However, our study sheds light on a novel neuronal function of Regase-1 on sleep and memory. Regnase-1 is an anti-inflammatory enzyme that inhibits mRNA translation during acute inflammatory responses. It localizes to the ribosomes on the surface of the endoplasmic reticulum (ER) and binds to translationally active mRNAs with specialized stem-loop structures at the 3′UTR (Uehata et al., 2013; Mino et al., 2015). When phosphorylated, Regnase-1 is released from the ER (Tanaka et al., 2019). Because functional Regnase-1 binds and degrades its bound mRNA, Regnase-1 inactivation leads to an increase of its target mRNA (Uehata et al., 2013). The target mRNAs of Regnase-1 in immune cells encode proinflammatory cytokines, which can then be expressed when Regnase-1 is inactivated. However, Regnase-1 has also been reported to modulate cytokines and neuronal injury in the microglia in rats (Liu et al., 2016). Regulation of Regnase-1 is usually rapid and transient, and its rapid response to microenvironmental changes, different pathological states and stress is critical for cellular adaption (Mao et al., 2017).

Earlier version: Our study reveals that the expression of Regnase-1 changes after training, and constitutive downregulation of Regnase-1 in ap α′/β′ neurons reduce sleep and causes deficits in learning and memory consolidation. We suggest that downregulation of Regnase-1 following training in fed flies may be crucial for the promotion of long-term memory, perhaps by promoting translation of specific transcripts. For instance, downregulation of Regnase-1 after training may release a pool of mRNAs that are normally targeted by it for decay and that are important for consolidation of memory. On the other hand, we find that constitutive loss of Regnase-1 impairs sleep-independent memory and also short-term memory, suggesting that it is required early in the learning and memory process, likely for learning. Interestingly, constitutive knockdown of Regnase-1 also reduces sleep and prevents sleep increase after training. However, our adult specific knockdown of Regnase-1 does not affect sleep, suggesting that expression of Regnase-1 during development is important for both memory and sleep. Given that ap α′/β′ neurons have a role in several aspects of memory, consolidation as well as short-term appetitive memory retrieval, how Regnase-1 acts within these neurons to regulate memory clearly requires more investigation.

Revised version: Our study reveals that the expression of Regnase-1 changes after training, and constitutive downregulation of Regnase-1 in ap α′/β′ neurons reduces sleep and causes deficits in learning and memory consolidation. Whether acute downregulation of Regnase-1 is necessary for sleep-dependent memory remains unclear, but it could play a role by promoting the translation of specific transcripts. For instance, post-training downregulation of Regnase-1 could release mRNAs that are usually targeted for decay but are critical for memory consolidation. Constitutive loss of Regnase-1 impairs sleep-independent memory and short-term memory, although the mechanisms are not known. We note that ap α′/β′ neurons have a role in several aspects of memory consolidation as well as short-term appetitive memory retrieval. Interestingly, constitutive knockdown of Regnase-1 also reduces sleep developmentally and prevents sleep increase after training in adult. While these could be independent effects, it is also possible that loss of Regnase-1 affects the development of the relevant MB neurons, thereby impacting all behaviors regulated by these neurons.

Earlier version: Our study of local molecular changes in ap α′/β′ neurons after training highlights how ribosomal RNA transcription and mRNA translation work in concert during the consolidation of sleep-dependent memory. Neurons are regulated to generate short bursts of boosted transcription and boosted translation, which are required for new protein synthesis. How sleep is involved in this boosted protein synthesis process is unclear, but we suggest that sleep affects several steps of the central dogma, and that deleterious consequences of sleep deprivation are mediated, at least in part, through effects on rRNA transcription and translation. Overall, our findings indicate the importance of maintaining RNA homeostasis for sleep and memory.

Revised version: Our study of local molecular changes in ap α′/β′ neurons after training suggests that ribosomal RNA transcription and mRNA translation might work in concert during the consolidation of sleep-dependent memory. How sleep is involved in this boosted protein synthesis process is unclear, but we suggest that RNA processing changes induce sleep, which promotes translation necessary for consolidation of long-term memory. However, these processes need to be teased apart experimentally. Overall, our findings demonstrate a role of RNA processing in sleep and memory, providing a foundation for future exploration of the mechanisms involved.

Added in the revised version: (A) PathwayPCA plot of all the differentially expressed genes (DEGs) across the three different conditions. The plot shows that genes responsible for PC1, which accounted for 35% of variance, largely encode proteins involved in transcription and biosynthesis, including RNA biosynthesis processes. (B) Top5 significant Gene Ontology (GO) biological process terms of all differentially expressed genes (DEGs) from three pairwise comparisons are shown. The x-axis represents the -log10 of the P-value, and the numbers on the bars indicate the count of genes identified for each GO term.

Added in the revised version: Effect of Polr1F and Regnase-1 knockdown on activity and sleep architecture.

Added in the revised version: (A-B) Total activity of flies is not altered by knockdown of Polr1F and Regnase-1. One-way ANOVA was used to calculate the p-values for each comparison. (C-D) Knockdown of Polr1F significantly increases the nighttime average sleep episode (SE) length in R35B12-Gal4>polr1F RNAi flies, while knockdown of Regnase-1 reduces it in R35B12-Gal4>regnase1 RNAi flies. The p-values for each comparison were calculated using the Kruskal–Wallis test with Dunn′s multiple comparisons test. (E-F) Knockdown of Polr1F significantly reduces the nighttime average sleep episode (SE) number in R35B12-Gal4>polr1FRNAi flies, but it is not significant compared to control groups in R35B12-Gal4>regnase1RNAi flies. The p-values for each comparison were calculated using the Kruskal–Wallis test with Dunn′s multiple comparisons test. ns = not significant, p > 0.05, *p < 0.05, **p < 0.01, ***p < 0.001, ****p < 0.0001.

Added in the revised version: Adult specific knockdown of Polr1F promotes sleep.

Added in the revised version: (A) Representative sleep traces of adult flies expressing Polr1F RNAi under the control of R35B12-Gal4/Tubulin-gal80ts. At permissive temperature, Gal80ts inhibits Gal4 activation of UAS; however, upon switching to restrictive temperature, Gal80ts is deactivated, leading to downstream activation of the UAS-regulated gene. (B) Sleep patterns were quantified over four consecutive days, with flies under permissive temperature on the first and last days and restrictive temperature on the second and third days. (C) Sleep quantification during the initial 3-hour period (ZT 0-3) after light onset for each day is shown. N = 46-48 individual flies per replicate with data from three independent replicates combined. The Kruskal-Wallis test with Dunn′s multiple comparisons test was used to compare three groups. ns = not significant, p > 0.05, *p < 0.05, **p < 0.01, ***p < 0.001, ****p < 0.0001.

Added in the revised version: Acute and chronic effects of pan-neuronal knockdown of Regnase-1 on sleep in adult flies.

Added in the revised version: (A) Schematic representation of the protocol for transient and chronic sleep measurements in nSyb-GS>regnase-1 RNAi flies. (B) Representative sleep traces plot of flies expressing Rengase-1 RNAi under the control of an inducible pan-neuronal driver (nSyb-GS>regnase-1 RNAi) with and without RU treatment (C) Quantification of sleep during the first 3 hours (ZT 9-12) after F1 progeny flies were loaded into RU- or RU+ DAM tubes at ZT8-T9. Sleep was measured starting at ZT9. N = 22-24 individual flies per replicate with data from two independent replicates combined. The Mann-Whitney test was used to compare RU+ group and RU-groups. (D) Representative average sleep traces of nSyb-GS>regnase-1 RNAi in the RU- and RU+ DAM tubes for three consecutive days. Chronic sleep effects of knockdown of Regnase-1 were measured based on sleep data from day 3 to day 5. (E) Quantification of average total sleep of nSyb-GS>regnase-1 RNAi and controls in the DAM tubes from (D). Unpaired t-test was used to compare between RU- and RU+ group. ns = not significant.

Added in the revised version: Adult specific knockdown of Regnase-1 has no effect on sleep.

Added in the revised version: (A) Representative sleep traces of adult flies expressing Regnase-1 RNAi under the control of R35B12-Gal4/Tubulin-gal80ts. (B) Sleep was quantified over four consecutive days, with flies under permissive temperature on the first and last days and restrictive temperature on the second and third days. N = 46 individual flies per replicate with data from three independent replicates combined. The Kruskal-Wallis test with Dunn′s multiple comparisons test was used to compare three groups. ns = not significant, p > 0.05, *p < 0.05, **p < 0.01.

Added in the revised version: Adult specific knockdown of Polr1F or Regnase-1 has no effect on peak activity.

Added in the revised version: (A) Maximal activity per minute for R35B12-Gal4; TubGal80ts>polr1F RNAi was quantified over four consecutive days, with flies under permissive temperature on the first and last days and restrictive temperature on the second and third days. (B) Maximal activity per minute for R35B12-Gal4; TubGal80ts>Regnase-1 RNAi was quantified. N = 40-48 individual flies per replicate with data from three independent replicates combined. The Kruskal-Wallis test with Dunn′s multiple comparisons test was used to compare three groups. ns = not significant, p > 0.05, *****p < 0.0001.

Added in the revised version: Regnase-1 overexpression in the ap α′/β′ neurons does not affect sleep.

Added in the revised version: (A) The efficacy of the RNAi used in the experiments was verified by real-time quantitative PCR. The different RNAi lines were crossed with nSyb-GS, and adult progeny were placed in either RU- or RU+ tubes overnight before being dissected for qPCR. Welch’s test was used for statistical analysis across all three RNAi, with *p<0.05 indicating significance. (B) Representative sleep traces and quantification of adult flies expressing Regnase-1 under the control of R35B12-Gal4. Sleep data were averaged over three consecutive days, with N = 31-32 individual flies per replicate, and combined data from two independent replicates were analyzed.

Added in the revised version: rRNA inhibitor (CX-5461) feeding does not affect sleep.

Added in the revised version: Feeding of the ribosome RNA inhibitor CX-5461 at a concentration of 0.2mM does not affect sleep in fed or starved flies. Statistical analysis shows that there is no significant difference between CX-5461-fed and control flies. The sample size is N=32 per group from two independent replicates. ns, not significant).

Removed in the revised version: Supporting information